# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 2  3  3
 2  4  2
 4  2  2
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 8.0
 8.0
 8.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 2  2  4
 3  4  2
 3  2  2
```

and we can get the transpose with

```julia
?transpose
```

Output:
```
search: [0m[1mt[22m[0m[1mr[22m[0m[1ma[22m[0m[1mn[22m[0m[1ms[22m[0m[1mp[22m[0m[1mo[22m[0m[1ms[22m[0m[1me[22m

[36m  transpose(A)[39m

  Lazy transpose. Mutating the returned object should appropriately mutate [36mA[39m.
  Often, but not always, yields [36mTranspose(A)[39m, where [36mTranspose[39m is a lazy
  transpose wrapper. Note that this operation is recursive.

  This operation is intended for linear algebra usage - for general data
  manipulation see [36mpermutedims[39m, which is non-recursive.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> A = [3+2im 9+2im; 8+7im  4+6im][39m
[36m  2×2 Matrix{Complex{Int64}}:[39m
[36m   3+2im  9+2im[39m
[36m   8+7im  4+6im[39m
[36m  [39m
[36m  julia> transpose(A)[39m
[36m  2×2 transpose(::Matrix{Complex{Int64}}) with eltype Complex{Int64}:[39m
[36m   3+2im  8+7im[39m
[36m   9+2im  4+6im[39m
```

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 2  2  4
 3  4  2
 3  2  2
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 24  22  18
 22  29  21
 18  21  17
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.518301   0.284985
 0.855072   0.227019
 0.0102318  0.365172
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
  3.5491061848783576
 21.762001096807644
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.0222053  0.0222053
 0.14772    0.14772
 0.563893   0.563893
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 8.625802335202868
 8.625802335202872
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.331598  0.220304  0.437562
 0.961089  0.110671  0.920583
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -0.9561604180565435
  3.2339835416675187
  1.28647325621642
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia
using LinearAlgebra
```

```julia
dot_v = dot(v, v)
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

```julia
dot_v = v ⋅ v
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
v'
```

Output:
```
1×3 adjoint(::Vector{Int64}) with eltype Int64:
 1  2  3
```

```julia
outer_v = v * v'
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

```julia
cross_v = cross(v, v)
```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

```julia

```